```python
import custom_funcs as cf
import numpy as np
import scipy as sp
from scipy.stats import pearsonr
```

# Problem Set on Machine Learning

## Question 1

Identify the original literature reference for the PhenoSense assay. Paste the URL to the PubMed article below, and write a one-paragraph summary on how the assay is done.

<!-- grade_id: phenosense, points: 3, grade: true, solution: true -->
## Your Answer Here

Double-click on this the answer cells, and start typing your answers. The cells will recognize Markdown. The most important things you may need are **bolded** text and *italicized* text, as well as:

1. Numbered
1. Lists
    1. With
    1. Sub-pointers

Also, you may like

- Bulleted 
- lists
    - with
    - subpointers
- indentations are 4 spaces in.

## Question 2

1. Identify the protease drugs for which the two drugs' resistance values are correlated. 
1. Write code below to calculate their correlation score. 
1. Search around the academic literature on these drugs, and provide an explanation as to why they would be correlated.

<!-- grade_id: drug_correlate, points: 3, grade: true, solution: true -->
**Answer to Question 1:**

The protease drugs that are most correlated are: **[put your answer here]**

```python
# This cell loads the data and cleans it for you, and log10 transforms the drug resistance values.
DATA_HANDLE = 'drug_data/hiv-protease-data.csv'  # specify the relative path to the protease drug resistance data
N_DATA = 8  # specify the number of columns in the CSV file that are drug resistance measurements.
CONSENSUS = 'sequences/hiv-protease-consensus.fasta'  # specify the relative path to the HIV protease consensus sequence
data, drug_cols, feat_cols = cf.read_data(DATA_HANDLE, N_DATA)
consensus_map = cf.read_consensus(CONSENSUS)
data = cf.clean_data(data, feat_cols, consensus_map)
for name in drug_cols:
    data[name] = data[name].apply(np.log10)
data.head()
```

Output:
```
FPV      ATV       IDV       LPV       NFV       SQV  TPV  DRV P1  \
SeqID                                                                           
2996   0.397940      NaN  1.212188       NaN  1.586587  1.206826  NaN  NaN  P   
4387  -0.154902      NaN -0.096910       NaN -0.096910  0.041393  NaN  NaN  P   
4432   0.176091      NaN  0.000000       NaN  0.342423  0.041393  NaN  NaN  P   
4482   0.591065      NaN  1.305351       NaN  1.334454  0.963788  NaN  NaN  P   
4486   0.977724  1.30103  0.913814  1.041393  1.857332  1.662758  NaN  NaN  P   

      P2 ... P90 P91 P92 P93 P94 P95 P96 P97 P98 P99  
SeqID    ...                                          
2996   Q ...   M   T   Q   L   G   C   T   L   N   F  
4387   Q ...   L   T   Q   I   G   C   T   L   N   F  
4432   Q ...   L   T   Q   I   G   C   T   L   N   F  
4482   Q ...   M   T   Q   L   G   C   T   L   N   F  
4486   Q ...   L   T   Q   I   G   C   T   L   N   F  

[5 rows x 107 columns]
```

<!-- grade_id: drug_correlation_score, points: 0, grade: false, solution: true -->
```python
"""
Complete the function below to compute the correlation score.

Use the scipy.stats.pearsonr(x, y) function to find the correlation score between two arrays of things. 
You do not need to type the whole name, as I have imported the pearsonr name for you, so you only have to do:

    pearsonr(x, y)

Procedure:

1. Select two columns' names to compare.
2. Make sure to drop NaN values. the pearsonr function cannot deal with NaN values. 
   (Refer to the Lecture notebook if you forgot how to do this.)
3. Pass the data in to pearsonr().

"""
def corr_score(drug1, drug2):
    ### BEGIN SOLUTION
    # Get the subset of data
    subset = data[[drug1, drug2]].dropna()
    # Return the pearsonr score.
    return pearsonr(subset[drug1], subset[drug2])
    ### END SOLUTION
```

<!-- grade_id: drug_correlation_tests, points: 1, grade: true, solution: false -->
```python
from itertools import combinations
assert corr_score('IDV', 'FPV') == (0.79921991532901282, 2.6346448659104859e-306)
assert corr_score('ATV', 'FPV') == (0.82009597442033089, 2.5199367322520278e-231)
assert corr_score('NFV', 'DRV') == (0.69148264851159791, 4.0640711263961111e-82)
assert corr_score('LPV', 'SQV') == (0.76682619729899326, 4.2705737581002648e-234)
```

<!-- grade_id: drug_correlation_mechanism, points: 3, grade: true, solution: true -->
**Answer to Question 3:**

# Problem 3

Write a block of code that plots the distribution of drug resistance across all protease proteins that are present.